## Assignment 2 - NLP using Rules
_Solutions have to be submitted by Monday, May 13, 23:59 (UTC+2)._

---

### Exercise 1: Finite-State Automaton and Finite-State Transducers (3 point)

**(a) Consider the task of finding spelling errors. Do you think a Finite State Transducer (FST) would be needed for this task or would a Finite State Automaton (FSA) suffice? Shortly explain your choice in one sentence.**

<!-- grade_id: cell-37695585c9f93b0e, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

**(b) What is the main difference between an FSA and an FST? Give an example of a task, not mentioned in the lecture, that an FST can do that an FSA can not.**

<!-- grade_id: cell-f78b09c1d51e73a9, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

**(c) Draw an FST that reads strings consisiting of the special characters ? and !, and rewrites the input string to have no more than two consecutive !'s or ?'s in them. Describe the FST using the notation from the lecture ($Q,Σ,q_0,F,\delta$).**

**Examples:**
- `!` $\to$ `!`
- `?` $\to$ `?`
- `!!` $\to$ `!!`
- `??` $\to$ `??`
- `!!!` $\to$ `!!`
- `???` $\to$ `??`
- `!?!` $\to$ `!?!`
- `!!?` $\to$ `!!?`
- `????????` $\to$ `??`
- `!!!!!!!!` $\to$ `!!`

_Note: You can include images in the markdown cell below by uploading an image file to the same directory as this notebook and then adding it in markdown via `![alt-text](image-file-name.png)`. You are free to choose between drawing the tree by hand (and then scanning it) or any digital drawing program of your liking. Make sure, however, that all texts and connections are clearly readable. If this is not the case, you will not get any points for this task._

<!-- grade_id: cell-623245b6e50cd851, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

### Exercise 2: Decision trees (3 points)

**Consider parsing chat messages in an online game where messages can be sent globally, only to the party, or only to one other user.**
- **To send a message globally the content just has to be typed: `<content>`.**
- **A message to the party is of the form `/party <content>`, sending `<content>` to the party.**
- **A message to a single user of the form `/whisper <user> <content>`.**
- **Additionally, there is a help command (`/help`) with which help can be displayed.**
- **Messages of a form not listed above are invalid and to be rejected.**

**(a) To simplify the later task of interpreting the message, we first want to tokenize the input using a decision tree. Draw this decision tree, according to the following rule: If and only if the first character is a slash, the slash is a token, else tokens are separated by spaces.**  

**This is similar to sentence splitting which was discussed in the lecture.** 

Hint: Characters are read in order one at a time. You can assume that the sentence is well formed (it does not start or end with a space and there are no double or triple spaces). If there is no next char, the remaining text is automatically considered to be a token.

Example: `'/party Have fun, good luck!'` $\to$ `['/', 'party', 'Have', 'fun,', 'good', 'luck!']`

_Note: You can include images in the markdown cell below by uploading an image file to the same directory as this notebook and then adding it in markdown via `![alt-text](image-file-name.png)`. You are free to choose between drawing the tree by hand (and then scanning it) or any digital drawing program of your liking. Make sure, however, that all texts and connections are clearly readable. If this is not the case, you will not get any points for this task._

<!-- grade_id: cell-04ba952973083ee1, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**(b) Now draw a decision tree in which the messages are categorized and dealt with sending the message to the right channel or displaying help as defined above. The starting point is the already tokenised message, e.g. `['/', 'party', 'Have', 'fun,', 'good', 'luck!']`.**

Hint: Tokens are read in order one at a time. You can assume there are functions to send remaining tokens to a specific channel or user and to print the help (e.g. "send remaining to user specified in next token").

<!-- grade_id: cell-5276b0a65ca3ee47, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

### Exercise 3: Lemmatization (1 point)

**Lemmatization can be done by looking up the corresponding lemma in a lexicon. This method has its limits though. Consider trying to lemmatize tokens in language with compound words (e.g. German). What limits do you see to the lexicon method? What method could be used to alleviate the problem?**  
Hint: [Compound -- Wikipedia](https://en.wikipedia.org/wiki/Compound_(linguistics)#Germanic_languages).<br>
Hint: Think about ad-hoc compounds.

<!-- grade_id: cell-5628d4c294e127ae, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

### Exercise 4: Template-based generation (3 points)

**Consider the task of template-based natural language generation in the domain of virtual assistants (e.g. Alexa, Siri, ...).**

**For a response to a prompt of your choice, come up with a template sentence that has to be filled with text from data. The template should have at least 3 slots to be filled. List and explain the following things:**

1. **The intended goal of the template and what it should convey.**
2. **The data (i.e. the entities, attributes and relations) that would be required to generate the sentence and the knowledge to order the retrieved data in the planning phase.**
3. **At least one rule that is applied in each, the Lexicalization and Referring expression generation step and why it makes sense to apply this rule.**
4. **The template itself (e.g. "At \<time> match, where \<entity-1>" ...).**
5. **An example of the template filled with information (e.g. "At this evening match, where THW Kiel" ...).**


_Note: The data and rules need to be different from those presented in the lecture._

<!-- grade_id: cell-6302a4abd6bd1700, points: 3, grade: true, solution: true -->


---

### Exercise 5: Programming task 1 (5 points)

**In this task you shall implement a rule-based tokenizer.**

**Write a function `my_tokenizer` that takes as an input a text of type string, parse the string character by character, and applies a set of rules in order to identify all positions where the text should be split. Then the function should return a list of strings representing the tokens extracted from the text. For examples of possible rules refer to the lecture part on rule-based approaches. You are free to add more rules that increase your performance.**

**Shortly describe your appraoch to tokenization in the docstring of the function.**

**Your code will be evaluated on a ground-truth sample of texts. The provided function `evaluate_tokenizer` will return the accuracy of your tokenizer on the sample. Try to optimize your rules to get the highest accuracy. The same ground-truth texts would be used to evaluate your code. Part of the credits for this task will depend on the achieved performance.**

Hint: You can have a look at the text_tokens.txt file to see what the ground truth tokens look like.

Example:`'guns should be banned because they are not needed in any domestic issue.'` $\to$ `['guns', 'should', 'be', 'banned', 'because', 'they', 'are', 'not', 'needed', 'in', 'any', 'domestic', 'issue', '.']`

_Note: An empty input should also return an empty list._

```python
def evaluate_tokenizer(tokenizer) -> float:
    """
    This function takes as an input a function `tokenizer` and returns the
    accuracy of this function evaluated against a text sample.

    The function `tokenizer` should accept a string as a parameter and
    return a list of strings.
    """
    # Read data files
    with open("text_raw.txt", "r", encoding="utf-8") as f:
        raw_texts = f.read().split("\n")
    with open("text_tokens.txt", "r", encoding="utf-8") as f:
        tokens = f.read().split("<new_line>")
    processed_texts = list(map(lambda x: x.strip().split("\n"), tokens))
    
    scores = []
    # Create token pairs from raw and processed texts
    for raw_text, text_tokens in zip(raw_texts, processed_texts):
        # Split text into tokens with provided splitter
        pred_tokens = tokenizer(raw_text)
        
        # Test for the amount of intersection between the two documents
        if len(pred_tokens) == 0:
            score = 0
        else:
            score = len(set(pred_tokens).intersection(set(text_tokens)))/len(set(text_tokens))
        
        scores.append(score)

    # Return the accuracy
    return round(sum(scores)/len(scores), 4)
```

<!-- grade_id: cell-dfc295e9b48904ff, points: 0, grade: false, solution: true -->
```python
DEBUG = True
def print_debug(msg:str):
    if DEBUG: print(msg)
    
def my_tokenizer(text: str) -> list[str]:
    """
    I've implemented a decision tree that iterates over the entire sentense and implements an algorithm simular to the one from the lecture(for sentence splitting).
    The function split() decides whether the string should be cut at the given index. If split returns an index (!= -1) then we append a new token to the tokens list.
    If split decides to not split (and returns -1) we check if we should split at the next index. 
    """
    tokens = list()
    start = 0
    cur = 0
    while cur < len(text)-1:
        end = split(text, cur) #last token index 
        
        if end == -1: #-1 means "dont split"
            cur += 1
            continue 
        
        tokens.append(text[start:end+1].replace(' ',''))
        
        start = end+1
        cur = start
    
    tokens.append(text[start : len(text)].replace(' ','')) #add the last token
    #tokens.append('<new_line>') #add end of paragraph token
    
    #remove whitespace tokens 
    tokens = list(filter(lambda token: token != "", tokens))
    
    for token in tokens:
        print_debug(token)
        
    return tokens

def split(text:str, cur:int):
    'implements my handcrafted rules'

    #split before whitespace
    if text[cur+1] in {' '}: 
        return cur
    
    #split before point only if there arent multiple points -> e.g. "..." is one token
    if text[cur] not in {'.'} and text[cur+1] in {'.'}: 
        return cur 
    # case if point has no whitespace afterwards
    if text[cur] in {"."} and text[cur+1].isalpha():
        return cur
    

    #split at other characters
    if text[cur+1] in {'(', ')', ',', ';', '?', '!','"',"-"}:
        return cur
    # case if other characters have no whitespace afterwards
    if text[cur] in {'(', ')', ',', ';', '?', '!','"',"-"}:
        return cur
    
    
    #case of n -> possible negations
    if text[cur+1] == "n":
        #case of n'... -> split before n
        if cur+3 < len(text) and text[cur+2] == "'" and text[cur+3].isalpha():
            return cur
        
        #case e.g. cannot ->split not of
        if cur+4 < len(text) and text[cur+2] == "o" and text[cur+3] == "t" and not text[cur+4].isalpha():
            return cur


    #handle "'" as sign for an abbreviation or quotation mark 
    if text[cur+1] in {"'"}:
        #case 1: white space before -> opening or closing quotation mark -> own token
        if text[cur] in {" "}:
            return cur+1
        
        #case 2: letter before and white space after -> closing quotation mark -> own token
        if text[cur].isalpha() and cur+2 < len(text) and  text[cur+2] in {" "}:
            return cur
       
        #case 2: letter before and after -> abbriviation e.g. n't -> "'" is no token
        if text[cur].isalpha() and cur+2 < len(text) and  text[cur+2].isalpha():
            # cut at "'" if there is no n before -> 's 're 'm 't 'll 've
            if text[cur] != 'n':
                return cur
            
        
    return -1

```

```python
accuracy = evaluate_tokenizer(my_tokenizer)
print(f"Accuracy of your tokenizer is: {accuracy}")
```

Output:
```
guns
should
be
banned
because
they
are
not
needed
in
any
domestic
issue
.
the
second
ammendment
was
put
in
place
because
of
fear
that
the
british
might
invade
america
again
or
take
control
of
the
government
.
if
this
were
the
case
the
people
would
need
weapons
to
defend
themselves
and
regain
america
.
the
british
are
n't
going
to
invade
so
we
do
n't
need
to
protect
our
selves
.
even
in
the
this
day
and
age
america
remains
increadible
safe
compared
to
many
other
nations
.
we
have
no
close
enemies
.
if
a
major
army
were
to
attack
us
a
few
men
with
pistols
or
shotguns
would
n't
do
much
against
a
soldier
with
an
ak47
or
tanks
or
bombers
.
guns
in
america
just
make
it
easier
for
crimes
to
be
committed
.
some
guns
should
never
be
considered
allowed
and
this
includes
all
semi
automatic
weapons
as
well
as
shotguns
.
poverty
,
drugs
,
and
lack
of
education
are
the
reasons
people
turn
to
guns
to
kill
.
guns
give
you
power
to
take
life
and
should
not
be
allowed
to
float
around
so
that
our
student
```

<!-- grade_id: cell-d2024eb9eec10d48, points: 5, grade: true, solution: false -->
```python
########## BEGIN OF AUTOMATED TESTS
# ----------------------------------------------------
# Tests for the `my_tokenizer` function
texts = [
    "text1 contains four tokens",
    "Text2 contains at six tokens. text1 contains four tokens"
]
texts_tokens = [my_tokenizer(text) for text in texts]
assert type(texts_tokens) == list
assert len(texts_tokens[0]) == 4
assert len(texts_tokens[1]) == 10
########## END OF AUTOMATED TESTS
```

Output:
```
text1
contains
four
tokens
Text2
contains
at
six
tokens
.
text1
contains
four
tokens
```

### Exercise 6: Programming task 2 (5 points)

**In this task you shall implement a rule-based morphological analyzer as a Finite-State Transducer (FST) for the task of English verb segmentation. Use the FST presented below (^ here is the symbol for a space (" ")):**


![fst](../../figures/a2_a6_fst.png)


**(a) Implement the function `verb_segmentation` that takes as an input a string and returns the verb segmentation of it as a list of strings (e.g: `'singing'` $\to$ `['sing', 'ing#']`) if it can be parsed using the defined state automata or `None` otherwise. The function should process the string character by character and try to build the segmented version of the verb.**

**We provide sample lexicons for regular verb stems, irregular verb stems and irregular past verb stems. Use these lexicons in your code to search for verb stem matches.**

**Use the docstring of the function to shortly describe your implementation approach.**

<!-- grade_id: cell-b03f7f62a4ff316a, points: 0, grade: false, solution: true -->
```python
# Regular verb stem Lexicon
reg_verb_stem = set(["walk", "fry", "talk"])
# Irregular verb stem Lexicon
irreg_verb_stem = set(["cut", "speak", "sing"])
# Irregular past verb Lexicon
irreg_past_verb = set(["caught", "ate", "eaten"])




def verb_segmentation(text: str) -> list[str]:
    """Describe implementation details here..."""
    
    state = "q0"
    stem = ""
    ending = ""
    i = 0

    '''
    while i < len(text):
        if state == "q0":
            #read new character 
            stem += text[i]
            print(stem)
            #compare current stem to verbs in lists
            if stem in irreg_verb_stem:
                state="q1"
            elif stem in reg_verb_stem:
                state="q2"
            elif stem in irreg_past_verb:
                state="q3"
                
        elif state == "q1":
            pass
        
        elif state == "q2":
            pass
        
        elif state == "q3":            
            pass
        '''
    #print(state)
    #i +=1
    return None

```

<!-- grade_id: cell-81e7baeb9a8b9ad6, points: 5, grade: true, solution: false -->
```python
########## BEGIN OF AUTOMATED TESTS
# ----------------------------------------------------
# Tests for the `nominal_number_inflection` function

# Regular verb stem Lexicon
reg_verb_stem = set(["walk", "fry", "talk"])
# Irregular verb stem Lexicon
irreg_verb_stem = set(["cut", "speak", "sing"])
# Irregular past verb Lexicon
irreg_past_verb = set(["caught", "ate", "eaten"])

verbs = ["person","walk", "walks", "walking", "caught", "cut", "cuts", "cutting"]
verb_segments = [verb_segmentation(verb) for verb in verbs]
assert verb_segments[0] == None
assert verb_segments[1] == ['walk', '#']
assert verb_segments[2] == ['walk', 's#']
assert verb_segments[3] == ['walk', 'ing#']
assert verb_segments[4] == ['caught', '#']
assert verb_segments[5] == ['cut', '#']
assert verb_segments[6] == ['cut', 's#']
assert verb_segments[7] == ['cut', 't', 'ing#']
########## END OF AUTOMATED TESTS
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[82], line 15
     13 verb_segments = [verb_segmentation(verb) for verb in verbs]
     14 assert verb_segments[0] == None
---> 15 assert verb_segments[1] == ['walk', '#']
     16 assert verb_segments[2] == ['walk', 's#']
     17 assert verb_segments[3] == ['walk', 'ing#']

AssertionError:
```

```python

```

Output:
```
Running cells with 'Python 3.12.2' requires the ipykernel package.

Run the following command to install 'ipykernel' into the Python environment. 

Command: '/opt/homebrew/bin/python3.12 -m pip install ipykernel -U --user --force-reinstall'
```